# Descriptive Stats

For the questions below, we're going to refer exclusively to an HHCAPS survey data set that is available in `/data/hhcaps.csv`.  Use whatever commands you want to calculate the information required to get to the answer.


Put your solution as the last line right before the test cell, using the form shown below for computing the number of rows in a dataframe.

```
answer = my_df.shape[0]
```

The assertions will often give you what the final answer should be, but you won't receive any points unless you compute the answer using code.  For instance, if you just typed `answer = 132` for the numer of rows instead of `answer = my_df.shape[0]`, you would not receive credit.


```python
%matplotlib inline
import pandas as pd

hhcaps = pd.read_csv('/data/hhcaps.csv')
hhcaps.head(3)
```

Output:
```
State  CMS Certification Number (CCN)*  \
0    AL                            17000   
1    AL                            17008   
2    AL                            17009   

                                       Provider Name  \
0  BUREAU OF HOME & COMMUNITY SERVICES           ...   
1  JEFFERSON COUNTY HOME CARE                    ...   
2  ALACARE HOME HEALTH & HOSPICE                 ...   

                                             Address                  City  \
0  201 MONROE STREET, THE RSA TOWER,  SUITE 1200 ...  MONTGOMERY             
1  2201 ARLINGTON AVENUE                         ...  BESSEMER               
2  2970 LORNA ROAD                               ...  BIRMINGHAM             

     Zip       Phone       Type of Ownership  Offers Nursing Care Services  \
0  36104  3342065341  Official Health Agency                          True   
1  35020  2059169500  Official Health Agency                          True   
2  35216  2058242680                   Local         
```

## #01 - 

How many columns does this file contain?

```python
hhcaps.shape
```

Output:
```
(12000, 39)
```

<!-- grade_id: q01-solution, points: 0, grade: false, solution: true -->
```python
answer = len(hhcaps.columns)
```

<!-- grade_id: q01-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 39)
```

## #02 - 

How many different values for State are there in this file?

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
answer = len(hhcaps['State'].unique())
```

<!-- grade_id: q02-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 55)
```

## #03 - 

Which of those State values has the highest frequency of occurence?

<!-- grade_id: q03-solution, points: 0, grade: false, solution: true -->
```python
hhcaps['State'].value_counts().unique
answer = hhcaps['State'].value_counts().idxmax()
```

<!-- grade_id: q03-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'TX')
```

## #04 - 

Which of those State values has the best average performance on the `Star Rating for health team communicated well with them` score?

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
filter_ans = hhcaps.groupby(['State','Star Rating for health team communicated well with them']).sum().reset_index()
filter_ans = filter_ans.groupby(['State']).mean().reset_index()
filter_ans.max()
answer = filter_ans[['State']][filter_ans['Star Rating for health team communicated well with them']== 4.5].reset_index()
answer = answer.iloc[0]['State']
```

<!-- grade_id: q04-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'ME')
```

## #05 - 

What was the average score on `Star Rating for how patients rated overall care from agency` for providers listed as having a `Type of Ownership` of `Hospital Based Program`

<!-- grade_id: q05-solution, points: 0, grade: false, solution: true -->
```python
filter_avg = hhcaps[hhcaps['Type of Ownership'] =='Hospital Based Program']
answer = filter_avg["Star Rating for how patients rated overall care from agency"].mean()
answer
```

Output:
```
3.6925207756232687
```

<!-- grade_id: q05-test, points: 2, grade: true, solution: false -->
```python
import numpy
assert(type(answer) == float or type(answer) == numpy.float64)
assert(round(answer,5) == round(3.6925207756232687,5))
```

## #06 - 

There is one column in the file that has the same value on every row.  What is the name of that column?

```python
filter_uniq = hhcaps.nunique()
answer = filter_uniq[filter_uniq == 1].index[0]
answer
```

Output:
```
'Offers Nursing Care Services'
```

<!-- grade_id: q06-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer.lower() == 'Offers Nursing Care Services'.lower())

```

## #07 - 

There are six (6) columns in the file that indicate (True or False) if the provider offers certain services.  Those column names all start with `Offers...`.  How many facilities offer **everything** that the survey was interested in asking about?  That is, how many rows have True in all six of those columns?

<!-- grade_id: q07-solution, points: 0, grade: false, solution: true -->
```python
filter_offers = hhcaps.columns[hhcaps.columns.str.startswith('Offers')]
filter_amount = hhcaps[filter_offers].sum(axis = 1)
answer = int((filter_amount == len(filter_offers)).sum())
answer
```

Output:
```
9029
```

<!-- grade_id: q07-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 9029)

```

## #08 - 

The survey results report `Number of completed Surveys` as one of the metrics.  What is the `median` of the number of completed surveys per facility?

Note that some rows don't have a valid number in them for the number of completed surveys.


<!-- grade_id: q08-solution, points: 0, grade: false, solution: true -->
```python
filter_surv = hhcaps['Number of completed Surveys'][hhcaps['Number of completed Surveys'] != 'Not Available']
answer = filter_surv.median()
answer
```

Output:
```
84.0
```

<!-- grade_id: q08-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == float or type(answer) == numpy.float64)
assert(answer == 84)

```

## #09 - 

How many providers are there from St. Louis, Missouri?

*Note that the City columnd may have trailing spaces that need to be accomodated or cleaned up*

<!-- grade_id: q09-solution, points: 0, grade: false, solution: true -->
```python
answer = hhcaps[(hhcaps['City'].str.strip() == 'SAINT LOUIS') & (hhcaps['State']=='MO')]['Provider Name'].nunique()
```

<!-- grade_id: q09-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == int)
assert(answer == 23)

```

# #10 - 

Which state got the highest percentage of 'top box' scorings, using the `Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)` field.

Ignore those rows where the data is **Not Available**

```python
field = 'Percent of patients who gave their home health agency a rating of 9 or 10 on a scale from 0 (lowest) to 10 (highest)'
ignore_rows = hhcaps[hhcaps[field] != 'Not Available'].reset_index()
ignore_rows[field] = ignore_rows[field].astype(float)
filter_high = ignore_rows.groupby('State')[field].mean().sort_values(ascending = False)
answer = filter_high.index[0]
```

<!-- grade_id: q10-test, points: 2, grade: true, solution: false -->
```python
assert(type(answer) == str)
assert(answer == 'MS')
```